# SIDAS502: Math Methods for Data Science 


# School of Information, University of Michigan


## Week2: Eigenvalue/Eigenvector/Linear Trasformation

### Version 1.2

This assignment is designed to demonstrate basic linear algebra and computational properties of matrices through identifying linearity across vectors, using row reduction to calculate eigenvalues and eigenvectors. Please read the directions carefully, as we want to avoid submissions that are marked incorrect due to formatting mistakes. You will be using sympy, numpy, and scipy for this assignment. 

Ensure that the following code is present at the beginning of your assignment:
- from sympy import *
- import numpy as np
- from numpy import linalg


Please enter your name: ""

## Part 1: Linear (In)dependence


<b>1.0</b> In this section, we will be determining if a set of vectors are linearly independent (if they are not, then they are linearly dependent). Below you will be asked to write a function to decide if 2, 3, or 4 inputed vectors in $\mathbf{R}^3$ are linearly independent.  


    


```python
from sympy import *
import numpy as np
from scipy import linalg
from numpy import linalg
```

<strong>1.1</strong> \[2 pt\] Impliment the function Indy2Vec that takes in two 3 dimentional vectors each represented as an array and tells whether they are linearly independent.


<!-- grade_id: cell-411ec698d9ae862d, points: 0, grade: false, solution: true -->
```python
def Indy2Vec(v1, v2):
    mat = np.array([v1, v2])
    _, indexes = Matrix(mat).T.rref()
    return True if len(indexes) == 2 else False

v1 = np.array([0, 5, 0])
v2 = np.array([0, -10, 0])
Indy2Vec(v1, v2)

```

Output:
```
False
```

<!-- grade_id: cell-45a28a1ab4ac003c, points: 2, grade: true, solution: false -->
```python
v1 = np.array([0, 5, 0])
v2 = np.array([0, -10, 0])
v3 = np.array([1, 2, 3])
v4 = np.array([-2, -4, -6])
assert Indy2Vec(v1, v2) == False, "Problem 1.1, Your code said that two linearly dependent vectors were independent"
assert Indy2Vec(v3, v4) == False, "Problem 1.1, Your code said that two linearly dependent vectors were independent"
assert Indy2Vec(v1, v3) == True, "Problem 1.1, Your code said that two linearly independent vectors were dependent"
assert Indy2Vec(v2, v3) == True, "Problem 1.1, Your code said that two linearly independent vectors were dependent"
```

<strong>1.2</strong> \[2 pt\] Impliment the function Indy2Vec that takes in three 3 dimentional vectors each represented as an array and tells whether they are linearly independent.

<!-- grade_id: cell-9437cf7256c37d72, points: 0, grade: false, solution: true -->
```python
def Indy3Vec(v1, v2, v3):
    # https://math.stackexchange.com/a/412566
    # http://sites.science.oregonstate.edu/math/home/programs/undergrad/CalculusQuestStudyGuides/vcalc/lindep/lindep.html#:~:text=For%20homogeneous%20systems%20this%20happens,has%20a%20non%2Dzero%20determinant.
    # simply find the determinant. if the det of M is not equal to 0, the vectors are linearly independent
    mat = Matrix(np.array([v1, v2, v3]))
    return True if mat.det() != 0 else False
    
v1=np.array([-1,3,4])
v2=np.array([6,-2,9])
v3=np.array([3,8,5]) 
Indy3Vec(v1, v2, v3)

```

Output:
```
True
```

<!-- grade_id: cell-158e70c4cabaca4d, points: 2, grade: true, solution: false -->
```python
v1=np.array([-1,3,4])
v2=np.array([6,-2,9])
v3=np.array([3,8,5]) 

assert Indy3Vec(v1, v2, v3) == True, "Problem 1.2, Your code said that two linearly independent vectors were dependent"
```

<strong>1.3</strong> \[1 pt\] Impliment the function Indy4Vec that takes in four 3 dimentional vectors each represented as an array and tells whether they are linearly independent.

<!-- grade_id: cell-d68d4ff3c19a01a9, points: 0, grade: false, solution: true -->
```python
def Indy4Vec(v1, v2, v3, v4):
    return False

v1=np.array([-1,3,4])
v2=np.array([6,-2,9])
v3=np.array([3,8,5])
v4=np.array([5,6,7])
Indy4Vec(v1, v2, v3, v4)

```

Output:
```
False
```

<!-- grade_id: cell-1b83b01be1cc2c58, points: 1, grade: true, solution: false -->
```python
v1=np.array([-1,3,4])
v2=np.array([6,-2,9])
v3=np.array([3,8,5])
v4=np.array([5,6,7])


assert Indy4Vec(v1, v2, v3, v4) == False, "Problem 1.3, Your code said that two linearly dependent vectors were independent"
```

<strong>1.4</strong> \[2 pt\] Impliment the function ThreeVec that takes in three 3 dimentional vectors each represented as an array,stacks the arrays vertically (each array is a row in a matrix), and returns a list in which the first element is the rank of the matrix, the second element is the rank of its transpose, the third element is its determinant rounded to two decimal places, the fourth element is the first element of its inverse rounded to two decimal places if its inverse exists and is the string "no inverse" if it does not exist.





<!-- grade_id: cell-d68d4ff3c19a01b8, points: 0, grade: false, solution: true -->
```python
def inverse(matrix):
    try:
        inv = np.linalg.inv(matrix)
        return round(inv[0][0], 2) if np.allclose(np.dot(matrix, inv), np.eye(3)) else "no inverse"
    except np.linalg.LinAlgError:
        return "no inverse"

def ThreeVec(v1, v2, v3):
    matrix = np.array([v1, v2, v3])
    matrix_rank = np.linalg.matrix_rank(matrix)
    transpose_rank = np.linalg.matrix_rank(matrix.T)
    det = np.linalg.det(matrix)
    inv = inverse(matrix)
    return [matrix_rank, transpose_rank, round(det, 2), inv]
    
v1=np.array([10,4.4,13])
v2=np.array([-5,-3.1,4])
v3=np.array([15,5.7,30])
ThreeVec(v1, v2, v3)

```

Output:
```
[2, 2, 0.0, 'no inverse']
```

<!-- grade_id: cell-1b83b01be1cc2c59, points: 2, grade: true, solution: false -->
```python
v1=np.array([10,4.4,13])
v2=np.array([-5,-3.1,4])
v3=np.array([15,5.7,30])



assert(ThreeVec(v1, v2, v3) == [2, 2, 0.0, 'no inverse']), "Problem 1.4, Your function does not compute the example correctly."










```

## Part 2: Eigenvalue/ Eigenvector


<strong>2.1</strong> \[1 pt\] What is the eigenvalue and eigenvectors of the matrix[[1,3],[2,5]]? Please store the eigenvalues in the variable called <strong>vals</strong>. The datatype for vals should be a numpy array.



<!-- grade_id: cell-571f59fab4ac61a1, points: 0, grade: false, solution: true -->
```python
v1= np.array([[1,3],[2,5]])
vals = np.linalg.eigvals(v1)
vals
```

Output:
```
array([-0.16227766,  6.16227766])
```

<!-- grade_id: cell-174c290441138ebc, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.1 are within this cell
```

<strong>2.2</strong> \[1 pt\] What are the eigenvectors of the matrix [[1,3],[2,5]]? Please store the eigenvectors in a variable called <strong>vecs</strong>. The datatype for vals should be an array.



<!-- grade_id: cell-ae7fca74e80f42d3, points: 0, grade: false, solution: true -->
```python
_, vecs = np.linalg.eig(v1)
vecs
```

Output:
```
array([[-0.93246475, -0.50245469],
       [ 0.36126098, -0.86460354]])
```

<!-- grade_id: cell-c9c31cdf9a7b77ad, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 2.2 are within this cell
```

<strong>2.3</strong> \[2 pt\] Find a 2 x 2 matrix with integer entries that has 10 and 25 as eigenvalues with eigen vectors $(1, 4)$ and $(1, -1)$ respectively.  Record it as an array in $ANS23$ with integer entries.  Hint: Put variables for entries in a 2 x 2 array, write out the equations that must be satisfied, and solve.

<!-- grade_id: cell-d68d4ff3c19a01b7, points: 0, grade: false, solution: true -->
```python
ANS23 = np.array([[22, -3], [-12, 13]])

```

<!-- grade_id: cell-1b83b01be1cc2c69, points: 2, grade: true, solution: false -->
```python
assert type(ANS23) == np.ndarray, "Problem 2.3, Your answer should be an array"
assert ANS23.shape == (2, 2), "Problem 2.3, Your array should be dimension 2 x 2"

assert np.isclose(ANS23 @ np.array([1, 4]), np.array([10, 40])).all(), "Problem 2.3, Your array should have [1, 4] as an eigenvector with eigenvalue 10"




```

## Part 3: Composition 


<strong>3.1</strong> \[1 pt\] Assume f(x)= 3x+6 and g(y)= 2y+1, what is the reduced equation of h(x)=g * f(x)=g(f(x))? Store your answer in the variable <strong>compo</strong> as a string type that is a valid python formula.  For example, "5 * x ** 2 + 4".  Note that neither "5 x ** 2 + 4" nor ""5 * x ^ 2 + 4" are valid python formulas.

<!-- grade_id: cell-81395b5bd2a197fb, points: 0, grade: false, solution: true -->
```python
compo="6 * x + 13"
```

<!-- grade_id: cell-f797c1459bf93874, points: 1, grade: true, solution: false -->
```python
assert type(compo) == str, "Problem 3.1, testing compo, type of value stored in variable does not match the expected type. Expecting String."
```

<strong>3.2</strong> \[1 pt\] Given x = 1, use the equation assigned to compo to solve. Store your answer in the variable <strong>compo1</strong> as an integer. 

<!-- grade_id: cell-beb471e3c613b1ec, points: 0, grade: false, solution: true -->
```python
compo1=19
```

<!-- grade_id: cell-850259d8a0f763ca, points: 1, grade: true, solution: false -->
```python
assert type(compo1) == int, "Problem 3.2, testing compo1, type of value stored in variable does not match the expected type. Expecting Integer."
```

## Part4: Linear Transformation 

<strong>4.1</strong> \[2 pt\] Assume we have an function: $f(x, y, z) = x$<sup>2</sup>$yz/3+x\sqrt{y}$+3z =10, is this function linear in x, y, or z?  For each variable, provide a string with a nonlinear term involving that variable or the empty string if none exists.  Again, report python executable strings.  For example, the equation might be written "(x ** 2) * y * z/3 + x * np.sqrt(y) + z + 3z = 10" in python

<!-- grade_id: cell-7cd69aa0c7535f5d, points: 0, grade: false, solution: true -->
```python
non_linear_X = "(x ** 2) * y * z/3"
non_linear_Y = "x * np.sqrt(y)"
non_linear_Z = ""
```

Output:
```
x**2*y/3 + 3
```

<!-- grade_id: cell-4e6661b95f6581b7, points: 1, grade: true, solution: false -->
```python
#hidden tests for problem 4.1 are within this cell
```